## Problem 3: How long distance individuals have travelled? (8 points)

In this problem the aim is to calculate the "distance" in meters that the individuals have travelled according the social media posts (Euclidean distances between points). In this problem, we will need the `userid` -column and the points created in the previous problem. You will need the shapefile `Kruger_posts.shp` generated in Problem 2 as input file.

Our goal is to answer these questions based on the input data:
- What was the shortest distance travelled in meters?
- What was the mean distance travelled in meters?
- What was the maximum distance travelled in meters?

**In your code, you should first:**
 - Import required modules.
 - Read in the shapefile as a geodataframe called `data`
 - Reproject the data from WGS84 projection into `EPSG:32735` -projection which stands for UTM Zone 35S (UTM zone for South Africa) to transform the data into the metric system.
 
*Store the result in a variable called `data`*!

<!-- grade_id: read_data, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import geopandas as gpd
import matplotlib.pyplot as plt
from shapely.geometry import Point, LineString
from pyproj import CRS
```

```python
# Read file from Shapefile
fp = "Kruger_posts.shp"
data = gpd.read_file(fp)
data = data.head(1000)
```

```python
# Let's make a backup copy of our data
#data_utm35s = data.copy()

# Reproject the data
data = data.to_crs(epsg=32735)
data.crs
```

Output:
```
<Derived Projected CRS: EPSG:32735>
Name: WGS 84 / UTM zone 35S
Axis Info [cartesian]:
- E[east]: Easting (metre)
- N[north]: Northing (metre)
Area of Use:
- name: Between 24°E and 30°E, southern hemisphere between 80°S and equator, onshore and offshore. Botswana. Burundi. Democratic Republic of the Congo (Zaire). Rwanda. South Africa. Tanzania. Uganda. Zambia. Zimbabwe.
- bounds: (24.0, -80.0, 30.0, 0.0)
Coordinate Operation:
- name: UTM zone 35S
- method: Transverse Mercator
Datum: World Geodetic System 1984 ensemble
- Ellipsoid: WGS 84
- Prime Meridian: Greenwich
```

- Check the crs of the input data. If this information is missing, set it as epsg:4326 (WGS84).
- Reproject the data from WGS84 to `EPSG:32735` -projection which stands for UTM Zone 35S (UTM zone for South Africa) to transform the data into metric system. (don't create a new variable, update the existing variable `data`!)"

<!-- grade_id: reproject, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
C:\Users\ALFOLD~1\AppData\Local\Temp/ipykernel_8200/962500874.py in <module>
      1 # REPLACE THE ERROR BELOW WITH YOUR OWN CODE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: problem_3_gdp_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# Check the data
print(data.head())
```

Output:
```
lat        lon         timestamp    userid  \
0 -24.980792  31.484633  2015-07-07 03:02  66487960   
1 -25.499225  31.508906  2015-07-07 03:18  65281761   
2 -24.342578  30.930866  2015-03-07 03:38  90916112   
3 -24.854614  31.519718  2015-10-07 05:04  37959089   
4 -24.921069  31.520836  2015-10-07 05:19  27793716   

                            geometry  
0  POINT (-4695752.719 14973674.275)  
1  POINT (-4748939.258 15014098.837)  
2  POINT (-4672729.591 14859391.193)  
3  POINT (-4679391.656 14969037.444)  
4  POINT (-4686373.982 14973910.589)
```

<!-- grade_id: problem_3_CRS_test, points: 1.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# Check that the crs is correct after re-projecting (should be epsg:32735)
print(data.crs)
```

Output:
```
epsg:32735
```

 - Group the data by userid

<!-- grade_id: group_data, points: 0, grade: false, solution: true -->
```python
grouped = data.groupby('userid')
```

<!-- grade_id: problem_3_grouped_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

#Check the number of groups:
assert len(grouped.groups) == data["userid"].nunique(), "Number of groups should match number of unique users!"
```

**Create LineString objects for each user connecting the points from oldest to latest:**

*Suggested steps:*
- Create an empty DataFrame called `movements`. 
- Create an empty column "geometry"
- Use a for-loop where you iterate over the grouped object. For each user's data: 
    - [sort](http://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.sort_values.html) the rows by timestamp 
    - create a LineString object based on the user's points
    - Add the LineString to the geometry column of the `movements` dataframe. You can also add the `userid` in a separate column (or use the userid as index).
- Convert `movements` into a `GeoDataFrame` (you can replace the DataFrame created in the previous steps with the GeoDataFrame). Remember to set the `geometry` column.
- Set the CRS of the ``movements`` GeoDataFrame as ``EPSG:32735`` 

<!-- grade_id: create_lines, points: 0, grade: false, solution: true -->
```python
movements = pd.DataFrame()
movements['userid'] = None
movements['geometry'] = None
movements = gpd.GeoDataFrame(movements, geometry='geometry', crs=CRS.from_epsg(32735).to_wkt())
for key, group in grouped:
    sort = group.sort_values('timestamp')['geometry'].to_list()
    if len(sort) > 1:
        ls = LineString(sort)
        dt = {'userid': key, 'geometry': ls}
        movements = movements.append(dt, ignore_index=True)
    else:
        ls = LineString([sort[0], sort[0]])
        dt = {'userid': key, 'geometry': ls}
        movements = movements.append(dt, ignore_index=True)
        
        

```

<!-- grade_id: problem_3_movements_tests, points: 3, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

#Check the result
print(type(movements))
print(movements.crs)
print(movements["geometry"].head())
```

Output:
```
<class 'geopandas.geodataframe.GeoDataFrame'>
PROJCRS["WGS 84 / UTM zone 35S",BASEGEOGCRS["WGS 84",ENSEMBLE["World Geodetic System 1984 ensemble",MEMBER["World Geodetic System 1984 (Transit)"],MEMBER["World Geodetic System 1984 (G730)"],MEMBER["World Geodetic System 1984 (G873)"],MEMBER["World Geodetic System 1984 (G1150)"],MEMBER["World Geodetic System 1984 (G1674)"],MEMBER["World Geodetic System 1984 (G1762)"],MEMBER["World Geodetic System 1984 (G2139)"],ELLIPSOID["WGS 84",6378137,298.257223563,LENGTHUNIT["metre",1]],ENSEMBLEACCURACY[2.0]],PRIMEM["Greenwich",0,ANGLEUNIT["degree",0.0174532925199433]],ID["EPSG",4326]],CONVERSION["UTM zone 35S",METHOD["Transverse Mercator",ID["EPSG",9807]],PARAMETER["Latitude of natural origin",0,ANGLEUNIT["degree",0.0174532925199433],ID["EPSG",8801]],PARAMETER["Longitude of natural origin",27,ANGLEUNIT["degree",0.0174532925199433],ID["EPSG",8802]],PARAMETER["Scale factor at natural origin",0.9996,SCALEUNIT["unity",1],ID["EPSG",8805]],PARAMETER["False e
```

**Finally:**
- Check once more the crs definition of your dataframe (should be epsg:32735, define the correct crs if this information is missing)
- Calculate the lenghts of the lines into a new column called ``distance`` in ``movements`` GeoDataFrame.

<!-- grade_id: calculate_distance, points: 0, grade: false, solution: true -->
```python
print(movements.crs)
movements['distance'] = movements['geometry'].length
```

Output:
```
PROJCRS["WGS 84 / UTM zone 35S",BASEGEOGCRS["WGS 84",ENSEMBLE["World Geodetic System 1984 ensemble",MEMBER["World Geodetic System 1984 (Transit)"],MEMBER["World Geodetic System 1984 (G730)"],MEMBER["World Geodetic System 1984 (G873)"],MEMBER["World Geodetic System 1984 (G1150)"],MEMBER["World Geodetic System 1984 (G1674)"],MEMBER["World Geodetic System 1984 (G1762)"],MEMBER["World Geodetic System 1984 (G2139)"],ELLIPSOID["WGS 84",6378137,298.257223563,LENGTHUNIT["metre",1]],ENSEMBLEACCURACY[2.0]],PRIMEM["Greenwich",0,ANGLEUNIT["degree",0.0174532925199433]],ID["EPSG",4326]],CONVERSION["UTM zone 35S",METHOD["Transverse Mercator",ID["EPSG",9807]],PARAMETER["Latitude of natural origin",0,ANGLEUNIT["degree",0.0174532925199433],ID["EPSG",8801]],PARAMETER["Longitude of natural origin",27,ANGLEUNIT["degree",0.0174532925199433],ID["EPSG",8802]],PARAMETER["Scale factor at natural origin",0.9996,SCALEUNIT["unity",1],ID["EPSG",8805]],PARAMETER["False easting",500000,LENGTHUNIT["metre",1],ID["EPSG"
```

<!-- grade_id: problem_3_distances_test, points: 2, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

#Check the output
movements.head()
```

Output:
```
userid                                           geometry      distance
0   234249  LINESTRING (-4769018.365 14944253.673, -476901...      0.000000
1   527064  LINESTRING (-4646195.657 14926347.125, -464626...  73383.339743
2   817930  LINESTRING (-4702693.831 15073343.240, -470269...      0.000000
3  2159344  LINESTRING (-4687383.025 14947706.969, -468578...  30443.513626
4  2265623  LINESTRING (-4673385.507 14862765.427, -467257...   3141.559996
```

You should now be able to print answers to the following questions: 

 - What was the shortest distance travelled in meters?
 - What was the mean distance travelled in meters?
 - What was the maximum distance travelled in meters?

```python
print(movements['distance'].min())
print(movements['distance'].mean())
print(movements['distance'].max())
```

Output:
```
0.0
34315.32176163399
1181752.6694053272
```

- Finally, save the movements of into a Shapefile called ``some_movements.shp``

<!-- grade_id: save_file, points: 0, grade: false, solution: true -->
```python
# Save the GeoDataFrame into a new Shapefile called 'polygon.shp'.
fp = 'some_movements.shp'

# Write the data into that Shapefile
movements.to_file(fp)
```

<!-- grade_id: problem_3_result_tests, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

import os

#Check if output file exists
assert os.path.isfile(fp), "Output file does not exits."
```

That's all for this week!